// intro1.rs
// Au sujet des commentaires `I AM NOT DONE` :
// Nous vous encourageons parfois à continuer à tester des choses sur un exercice
// donné, même après que vous l’ayez résolu. Si tout fonctionne et que vous sentez
// qu’il est temps de passer à l'exercice suivant, supprimez le commentaire
// `I AM NOT DONE` ci-dessous (littéralement "JE N'AI PAS FINI").
// Exécutez `rustlings hint intro1` pour avoir un conseil.

// I AM NOT DONE

fn main() {
    println!("Bonjour et");
    println!(r#"       bienvenue dans...                  "#);
    println!(r#"                 _   _ _                  "#);
    println!(r#"  _ __ _   _ ___| |_| (_)_ __   __ _ ___  "#);
    println!(r#" | '__| | | / __| __| | | '_ \ / _` / __| "#);
    println!(r#" | |  | |_| \__ \ |_| | | | | | (_| \__ \ "#);
    println!(r#" |_|   \__,_|___/\__|_|_|_| |_|\__, |___/ "#);
    println!(r#"                               |___/      "#);
    println!(r#"                                          "#);
    println!(r#"                          ...en français !"#);
    println!();
    println!("Cet exercice compile sans problème. Les autres exercices contiennent une erreur de");
    println!("compilation ou de logique. Le concept de base derrière Rustlings est de corriger");
    println!("ces erreurs et de résoudre les exercices. Bonne chance !");
}
